# COMP SCI 1015 IAP - W04 - Practical

## Task 1 - Price range (30 marks)

Assume we you are purchasing a laptop for the school. The school has decided the `lower` and `upper` limits of the purchase.

Please implement a function `purchase()` that takes three parameters `lower`, `upper` and `laptop_cost` and decide wheather the proposed laptop is `too expensive`, `too cheap`, or `acceptable`.

The function should return one of the following strings: `too cheap`, `too expensive`, or `acceptable`.

```python
def purchase(lower, upper, laptop_cost):
    # INSERT YOUR CODE BELOW
    if (laptop_cost > upper):
        return("too expensive")
    elif (laptop_cost < lower):
        return("Too cheap")
    else:
        return("acceptable")

# ================
# TEST CASES BELOW
s = purchase(500, 1000, 800) # return 'acceptable'
print(s)
s = purchase(1000, 1500, 2000) # return 'too expensive'
print(s)
```

Output:
```
acceptable
too expensive
```

## Task 2 - Estimate Pi (30 marks)

The Indian Mathematician [Srinivasa Ramanujan](https://en.wikipedia.org/wiki/Srinivasa_Ramanujan) discovered that the following infinite series could be used to approximate $1/\pi$. His series is the basis of one of the fastest algorithms for approximating $\pi$.

$\frac{2\sqrt{2}}{9801}\sum_{k=0}^{\infty}\frac{(4k)!(1103+26390k)}{(k!)^{4}396^{4k}}=\frac{1}{\pi}$

Please complete the implementation of the `estimate_pi(n)`. As computing $\sum_{k=0}^{\infty}$ would take infinite time, instead you will be looking at: $\sum_{k=0}^{n}$

```python
import math

def estimate_pi(n):
    '''print the estimated pi value '''
    s = 0
    k = 0
    # INSERT YOUR CODE BELOW
    # ~ 6 lines
    while k <= n:
        term = (math.factorial(4*k)*(1103+26390*k)) / ((math.factorial(k)**4) * 396**(4*k))

        s += term
        k+=1
    # END OF YOUR CODE
    return 1/ (2*math.sqrt(2)/9801 * s)


# ================
# TEST CASES BELOW
est = estimate_pi(0) # est_pi: 3.141592730013306 pi: 3.141592653589793     
print(f'est_pi: {est:.15f} pi: {math.pi:.15f}')
est = estimate_pi(3) # est_pi: 3.3.141592653589793 pi: 3.141592653589793     
print(f'est_pi: {est:.15f} pi: {math.pi:.15f}')
```

Output:
```
est_pi: 3.141592730013306 pi: 3.141592653589793
est_pi: 3.141592653589793 pi: 3.141592653589793
```

## Task 3: Three dice game (40 marks)

Please write a simple three dice program that keep rolling three dices until these three dices have the same value. The program then reports how many times it roll the dice. This activity tests what you have learned so far and might be challenging. Let TA know if you need help!

```python
import random

def roll():    
    '''return 1 if all dices have same values. return 0 otherwise.'''
    dice1 = random.randint(1, 6) 
    dice2 = random.randint(1, 6)
    dice3 = random.randint(1, 6)
    print(f'roll: {dice1} {dice2} {dice3}')
    if (dice1 == dice2 and dice2 == dice3):
        return 1
    else:
        return 0
    
def play():
    '''play the game, keep rolling until three dices have same value'''
    # YOUR CODE BELOW
    # ~ 5 LINES
    counter = 0
    result = False
    while (result != True):
        result = roll()
        counter+=1
    print(f'The player rolled {counter} times!')
    return counter

play()
```

Output:
```
roll: 1 2 4
roll: 4 1 4
roll: 1 6 5
roll: 6 2 2
roll: 5 3 2
roll: 2 4 5
roll: 1 3 4
roll: 2 3 4
roll: 4 1 6
roll: 6 5 5
roll: 5 5 6
roll: 4 3 2
roll: 2 1 5
roll: 2 4 1
roll: 5 3 1
roll: 2 1 6
roll: 3 2 6
roll: 6 5 3
roll: 5 3 1
roll: 2 3 4
roll: 4 3 6
roll: 1 1 4
roll: 4 5 2
roll: 4 1 1
roll: 4 5 4
roll: 3 5 1
roll: 4 2 3
roll: 3 4 2
roll: 5 1 4
roll: 3 5 3
roll: 3 4 4
roll: 5 5 3
roll: 5 2 1
roll: 6 4 2
roll: 2 5 1
roll: 5 5 1
roll: 6 3 1
roll: 4 6 2
roll: 4 3 2
roll: 5 4 1
roll: 1 6 4
roll: 4 5 2
roll: 3 1 6
roll: 2 4 3
roll: 3 4 5
roll: 1 2 2
roll: 5 4 2
roll: 4 3 4
roll: 6 6 5
roll: 6 4 3
roll: 1 1 1
The player rolled 51 times!
51
```

Tests for Task1 (points: 30)
```
>>> purchase(500, 1000, 800) == 'acceptable'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> purchase(1000, 1500, 2000) == 'too expensive'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> purchase(800, 1000, 700) == 'too cheap'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task2 (points: 30)
```
>>> estimate_pi(0) == 3.141592730013306
False
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> estimate_pi(5) == 3.141592653589793
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task3 (points: 40)
```
>>> play() == 47
roll: 2 4 1
roll: 2 1 2
roll: 6 4 3
roll: 6 1 4
roll: 2 4 6
roll: 1 5 2
roll: 4 5 2
roll: 3 1 3
roll: 2 4 5
roll: 2 1 4
roll: 2 5 3
roll: 5 6 1
roll: 5 2 4
roll: 3 2 5
roll: 5 2 5
roll: 1 5 4
roll: 4 3 5
roll: 2 6 1
roll: 6 2 1
roll: 6 1 6
roll: 4 1 5
roll: 3 2 2
roll: 4 6 4
roll: 4 1 5
roll: 4 4 1
roll: 6 1 2
roll: 4 1 5
roll: 6 2 4
roll: 1 4 1
roll: 6 2 6
roll: 4 5 3
roll: 2 1 3
roll: 4 4 3
roll: 6 1 5
roll: 3 5 5
roll: 2 6 1
roll: 5 2 6
roll: 6 2 3
roll: 1 4 4
roll: 6 6 4
roll: 4 5 1
roll: 1 4 1
roll: 5 5 3
roll: 1 2 4
roll: 3 5 2
roll: 6 3 2
roll: 2 6 1
roll: 3 5 3
roll: 4 3 2
roll: 6 5 4
roll: 6 5 3
roll: 2 4 1
roll: 2 2 2
The player rolled 52 times!
False
```
failure_message: Incorrect answer, try again!
success_message: Good job!
